# Grundlagen der Künstlichen Intelligenz - Wintersemester 2025/26


# Übung 1: Einführung

---

> **Grundlagen der künstlichen Intelligenz** im Wintersemester 2025/2026
>
> - Prof. Pascal Friederich, pascal.friederich@kit.edu
> - T.T.-Prof. Peer Nowack, peer.nowack@kit.edu
> 
> Übungsleiter für die aktuelle Übung: Henrik Schopmans, henrik.schopmans@kit.edu
> 
> ⚠️ Bei allgemeinen Fragen zu Aufgabenstellungen etc. bitte im Ilias Forum posten, sodass wir diese Fragen nur ein Mal für alle Studierenden beantworten können. E-Mail an den Übungsleiter bitte nur für individuelle Fragen, deren Beantwortung nicht für alle Studierenden relevant ist.

---

**Wichtig**: Dieses erste Übungsblatt wird nicht gewertet und soll euch helfen, euch mit Python, Jupyter, Pandas und Numpy vertraut zu machen. Außerdem bietet es euch
die Möglichkeit, technische Probleme zu erkennen.

- **Wir nutzen ein automatisiertes Grading System zur Bewertung der Jupyter Notebooks (siehe unten für Details)**
- **Bitte ladet dieses Übungsblatt daher bei Ilias als Abgabe hoch, auch wenn ihr es nur teilweise (oder gar nicht!) bearbeitet habt.**
- **Auf diese Weise könnt ihr technische Probleme beim Grading finden und wir können sie gemeinsam in der ersten Saalübung lösen.**
- **Ab dem zweiten Übungsblatt erwarten wir, dass die technischen Probleme gelöst sind.**
- **Nutzt daher bitte diese Möglichkeit!**

### Übungsteam

- Mozhgan Amiramjadi, mozhgan.amjadi@kit.edu
- Lina Rennstich, lina.rennstich@kit.edu
- Ulrich Oberhofer, ulrich.oberhofer@kit.edu
- Laura Ruple, laura.ruple@kit.edu
- Henrik Schopmans, henrik.schopmans@kit.edu
- Jonas Teufel, jonas.teufel@kit.edu

# Gruppenabgabe

Die Übungsblätter können in Gruppen von bis zu **3 Studierenden** abgegeben werden. **Jede Person aus der Gruppe muss die finale Version der Abgabe über Ilias hochladen**, es genügt nicht, dass nur eine Person aus der Gruppe dies tut. Es ist prinzipiell möglich, im Laufe des Semesters sich einer neuen Gruppe anzuschließen, sollte sich die eigene Gruppe vorzeitig auflösen. Jede Gruppe muss ihre eigene Lösung hochladen, wir werden die Abgaben auf Duplikate überprüfen.

Die Gruppen werden automatisch erfasst, **gebt deshalb die u-Kürzel eurer Gruppenmitglieder in die folgende Zelle ein.** Falls eure Gruppe nur aus 2 Studierenden besteht, oder ihr alleine abgebt, lasst die verbleibenden Felder frei. Hier ein **Beispiel** für eine Gruppe bestehend aus uabcd und uefgh:

_U-Kürzel der Gruppenmitglieder:_

_Mitglied 1: uabcd_

_Mitglied 2: uefgh_

_Mitglied 3:_


U-Kürzel der Gruppenmitglieder:

Mitglied 1:

Mitglied 2:

Mitglied 3:


# Auto-grading

Wir nutzen ein auto-grading System, welches eure abgegebenen Jupyter Notebooks automatisch analysiert und über
hidden tests auf Richtigkeit prüft. Über diese Tests werden die Punkte bestimmt, die ihr für das Übungsblatt erhaltet.

Damit das auto-grading reibungslos funktioniert bitte folgende Dinge beachten:

- Notebook muss den Dateinamen "ex_01_introduction.ipynb" haben
- PDF und Jupyter Notebook einzeln im Ilias hochladen (nicht als Zip!) 
    - Achtung: Für Übung 1 existiert kein handschriftlicher Teil.
- Vor dem Abgeben eines Notebooks bitte testen, dass alles von vorne bis hinten ohne Fehler durchläuft.
- Zellen, welche mit "##### DO NOT CHANGE #####" markiert sind dürfen nicht bearbeitet oder gelöscht werden
- Eure Lösung muss in die richtige Zelle (markiert mit "# YOUR CODE HERE") eingetragen werden.
  - Dabei natürlich den NotImplementedError löschen!
- Generell bitte **keine Zellen löschen** und auch **keine Zellen hinzufügen**. Die Zellen, in die eure Lösung eingetragen werden soll, sind bereits vorhanden (markiert mit "# YOUR CODE HERE").
- Es gibt potentiell scheinbar leere Zellen, die auch mit "" markiert sind. Auch diese dürfen nicht bearbeitet oder gelöscht werden.
  - Falls dies doch gemacht wird, dann wird das automatische Grading nicht funktionieren und ihr erhaltet keine Punkte.
  - Wir werden hier strikt handeln und keine Ausnahmen machen, falls jemand doch Zellen verändert, die eindeutig als readonly markiert sind!
- Die Jupyter Notebooks haben inline Tests (für euch sichtbar), welche euer Ergebnis auf grobe Richtigkeit überprüfen.
  - Diese sind primär für euch, um Fehler zu erkennen und zu korrigieren.
  - Die inline Tests, die ihr im Notebook sehen könnt, sind allerdings nicht die Tests welche für das Grading verwendet werden!
  - Die inline Tests sind eine notwendige aber keine hinreichende Bedingung, um beim Grading der Aufgabe Punkte zu erhalten!


## Einleitung
Wenn du mit Python nicht vertraut bist, empfehlen wir eines der folgenden Online-Tutorials:

* https://learnxinyminutes.com/docs/python/
* http://www.datacamp.com/community/tutorials/tutorial-jupyter-notebook
* https://www.learnpython.org/
* https://www.w3schools.com/python/
* https://automatetheboringstuff.com/

## Coronapandemie (nicht bewertet)

Um typische Datascience-Workflows in Python kennenzulernen analysieren wir hier kurz Falldaten zur Corona-Pandemie. Dazu laden wir zunächst einen historischen Datensatz über die Corona
Infektionen weltweit vom Europäischen Zentrum für die Prävention und
Kontrolle ([link](https://www.ecdc.europa.eu/sites/default/files/documents/COVID-19-geographic-disbtribution-worldwide.xlsx)) herunter.
Das können wir mit dem Python Package ``requests`` machen:

```python
##### DO NOT CHANGE #####
import os
from datetime import datetime
import matplotlib.dates as mdates
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import requests
import scipy.optimize

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
data_url = "https://www.ecdc.europa.eu/sites/default/files/documents/COVID-19-geographic-disbtribution-worldwide.xlsx"
data_file = "COVID-19-geographic-disbtribution-worldwide.xlsx"
if not os.path.exists(data_file):
    print("Downloading dataset ...")
    with open(data_file, "wb") as f:
        f.write(requests.get(data_url).content)
    print("Downloading dataset done.")

##### DO NOT CHANGE #####
```

Nun laden wir den Datensatz mit ``pandas``, welches ein ``DataFrame`` Objekt erstellt. Wir schauen uns die ersten 5 Zeilen an:

```python
##### DO NOT CHANGE #####
corona_world = pd.read_excel(data_file)

##### DO NOT CHANGE #####
```

```python
corona_world.head()
```

Das ``DataFrame`` Objekt erlaubt Zugriff auf die enthaltenen Daten über die Methoden .iloc (index-basierte Selektion) und .loc (name-basierte Selektion):

```python
print(corona_world.index)
print(corona_world.columns)

print(corona_world.iloc[0:5,[0,5]])
print(corona_world.loc[0:5,["month"]])
```

Nun selektieren wir nur die Einträge welche zu Deutschland gehören. Mit ``corona_countries=="Germany"`` erhalten wir eine boolean Mask, welche wir auf `corona_world` anwenden können. Weiterhin selektieren wir nur columns, welche für uns relevant sind. Außerdem flippen wir die finalen numpy arrays, sodass diese von der Vergangenheit forwärts laufen.

```python
##### DO NOT CHANGE #####
# Select data
corona_countries = corona_world["countriesAndTerritories"]
print(corona_countries == "Germany")
corona_germany = corona_world[corona_countries == "Germany"]
corona_germany = corona_germany[[
    "dateRep", "cases", "deaths",
    "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000"
]]
time_germany = corona_germany["dateRep"].values
cc_germany = corona_germany[
    "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000"
].to_numpy()
time_germany = np.flip(time_germany)
cc_germany = np.flip(cc_germany)

##### DO NOT CHANGE #####
```

Wir visualisieren nun die 14-Tage-Inzidenz für Deutschland (`Cumulative_number_for_14_days_of_COVID-19_cases_per_100000`). Hierfür nutzen wir das Package ``matplotlib.pyplot``, welches wir in zukünftigen Aufgabenblättern noch oft nutzen werden. [Hier](https://matplotlib.org/stable/tutorials/introductory/usage.html#sphx-glr-tutorials-introductory-usage-py) findet sich ein guter Guide zu diesem Package.

```python
##### DO NOT CHANGE #####
# Date functionality
months = mdates.MonthLocator()  # every month

##### DO NOT CHANGE #####
```

```python
fig, ax = plt.subplots()  # generate a new plot
ax.plot(time_germany, cc_germany)  # plot data
ax.xaxis.set_major_locator(months)  # modify axis
fig.autofmt_xdate()
plt.grid(True, "both")
plt.title("Germany")
plt.ylabel("14 days cases per 100k")
plt.xlabel("Time [year-month]")
plt.show()
```

Nun folgen ein paar Aufgaben, die du selbst beantworten sollst.
Hierfür kannst du entweder das pandas dataframe und seine Operationen direkt verwenden,
oder die Daten konvertieren und mit `numpy` arbeiten. Versuche wenn möglich immer eingebaute
Funktionalitäten von `numpy` oder `pandas` zu nutzen, da diese deutlich schneller sind 
als pures Python.

Dokumentation findet sich hier: [numpy](https://numpy.org/doc/) und
[pandas](https://pandas.pydata.org/docs/getting_started/index.html#getting-started)

**1.** Liste die Zahl der Todesfälle in Verbindung mit Covid-19 relativ zur mittleren Bevölkerungszahl (gemittelt über alle Einträge für das jeweilige Land) für jedes Land. Welches Land hat die höchste Todesrate? Welches die niedrigste?

Für diese Aufgabe bietet es sich an, die Ergebnisse in Dictionaries mit einem Key für jedes Land zu speichern. Du kannst dir weiterhin folgende möglicherweise hilfreichen Funktionen anschauen:

[pandas](https://pandas.pydata.org/docs/getting_started/index.html#getting-started):
``.unique``, ``.groupby``, ``.mean``, ``.sum``, ``.sort_values``, ``.get_group`` and ``.from_dict``:

<!-- grade_id: exercise_1, points: 0, grade: false, solution: true -->
```python
# Fülle die folgenden Objekte:

countries = [] # Liste aller Länder (jedes Land nur ein Mal)
population = {} # Durchschnitt der Bevölkerungszahl über alle Einträge für jedes Land
deaths = {} # Summe aller Todesfälle für jedes Land
rel_deaths = {} # Summe aller Todesfälle geteilt durch Bevölkerungszahl für jedes Land
country_max_mortality = "" # Welches Land hatte die höchste relative Todesrate?

# YOUR CODE HERE
raise NotImplementedError()
```

**2.** Nun möchten wir die kumulative Anzahl der Fälle (summierte Anzahl der Fälle bis zu einem Datum), die kumulative Anzahl der verstorbenen Fälle und die kumulative Anzahl der genesenen Fälle für Deutschland über der Zeit plotten.

Stelle sicher, dass die Daten die richtige Reihenfolge für eine kumulative Summe haben.

Tipp: Schau dir die Funktion ``np.cumsum`` an.

Tipp zum Bestimmen der kumulative Anzahl der genesenen Fälle:
Wir nehmen an, dass man nach 14 Tagen nach der Infektion / positivem Test entweder genesen oder verstorben ist.
Daher berechnen wir die kumulative Anzahl der genesenen Fälle, indem wir die Anzahl der kumulativen Fälle vor vierzehn Tagen minus der aktuellen kumulativen Anzahl der Todesfälle rechnen.

<!-- grade_id: exercise_2, points: 0, grade: false, solution: true -->
```python
cum_death = np.array([])
cum_cases = np.array([])
cum_recov = np.array([])

# YOUR CODE HERE
raise NotImplementedError()
```

Wir wollen nun analysieren, wie gut man mit einfachen Modellen den Verlauf der 
Pandemie im Spätjahr 2020 (nach 01.10.2020) in Deutschland vorhersagen kann.

**3.** Zunächst selektieren wir die Daten vom 01.08.2020 bis 01.10.2020 zum Fitten unserer Modelle. Beachte, dass wir hier numpy arrays vom Typ ``dtype="datetime64[ns]"`` verwenden. 

Wir müssen die passende Zeitperiode selektieren und dann durch ``np.timedelta64(1,"D")``, was dem Zeitraum eines Tages entspricht, teilen. Dann erhalten wir ein array, welches die Tage vom 01.08.2020 aufwärts zählt.

Als Testdaten selektieren wir alle Daten ab dem 01.08.2020, also auch über den 01.10.2020 hinaus.

```python
date_start = datetime(2020, 8, 1)
date_stop = datetime(2020, 10, 1)
selection = np.logical_and(time_germany >= np.datetime64(date_start),
                           time_germany < np.datetime64(date_stop))

# Selektiere Daten
x_cases = np.array(time_germany[selection] - np.datetime64(date_start)) / np.timedelta64(1, "D")
y_cases = np.array(cc_germany)[selection]
x_test = time_germany[time_germany >= np.datetime64(date_start)]
x_test_days = (x_test - np.datetime64(date_start)) / np.timedelta64(1, "D")
```

Nun versuchen wir ein paar analytische Funktionen an die Daten zu fitten, wofür wir ``scipy.optimize.curve_fit`` verwenden. Falls der Fit nicht auf Anhieb konvergiert, versuche die bounds und den initial guess von Hand zu ändern.

Versuche die folgenden Ausdrücke. Implementiere diese zunächst mit der Hilfe von Funktionen wie ``np.square``, ``np.exp`` oder ``np.power`` etc.

* $f_1(x)= ax+b$
* $f_2(x)= c e^{a (x-b)}$
* $f_3(x)= a x^2 + b x + c$
* $f_4(x)= a x^3 + b x^2 + cx + d$
* $\dots$

<!-- grade_id: exercise_3_functions, points: 0, grade: false, solution: true -->
```python
# Funktionen, welche implementiert werden sollen:
def fun_1(x, a, b):
    return a * x + b


def fun_2(x, a, b, c):
    pass


def fun_3(x, a, b, c):
    pass


def fun_4(x, a, b, c, d):
    pass


# Du kannst auch eigene zusätzliche Ideen implementieren...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise_3_fit, points: 0, grade: false, solution: true -->
```python
param_1, _ = scipy.optimize.curve_fit(fun_1, x_cases, y_cases)
param_2, _ = None, None  # fit fun_2
param_3, _ = None, None  # fit fun_3
param_4, _ = None, None  # fit fun_4

# YOUR CODE HERE
raise NotImplementedError()

y_fit_1 = fun_1(x_test_days, *param_1)
y_fit_2 = None  # evalute fit 2
y_fit_3 = None  # evalute fit 3
y_fit_4 = None  # evalute fit 4

# YOUR CODE HERE
raise NotImplementedError()

fig, ax = plt.subplots()  # generate a new plot
ax.plot(time_germany, cc_germany)  # plot data
ax.plot(x_test, y_fit_1, label="f1")  # plot data

# YOUR CODE HERE
raise NotImplementedError()

ax.xaxis.set_major_locator(months)  # modify axis
fig.autofmt_xdate()
plt.grid(True, "both")
plt.ylim([-20, 370])

plt.legend()
plt.title("Germany")
plt.ylabel("14 days cases per 100k")
plt.xlabel("Time [year-month]")

plt.show()  # make plot
```

- (1) Welche Funktion extrapoliert am besten über den 01.10.2020 bis zum Ende des Jahres?
- (2) Ändert sich das, wenn du nur an die letzten 14 Tage vom 17.09.2020 bis zum 01.10.2020 fittest?

<!-- grade_id: exercise_3_best_function, points: 0, grade: false, solution: true -->
```python
answer_best_fun = None  # Beste Funktion für die erste Frage (int)?
answer_best_fun_14_days = None  # Welche Funktion ist im zweiten Fall (2) am besten? (int)

# YOUR CODE HERE
raise NotImplementedError()
```

**Das war es auch schon für das erste Übungsblatt.**

**Nun folgen noch ein paar Checks für dich, mit welchen du deine Ergebnis grob überprüfen kannst:**

<!-- grade_id: correct_exercise_1, points: 6, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: correct_exercise_1 - possible points: 6

assert isinstance(countries, list) or isinstance(countries, np.ndarray)
assert isinstance(rel_deaths, dict)
assert isinstance(population, dict)
assert isinstance(deaths, dict)
assert isinstance(country_max_mortality, str)


##### DO NOT CHANGE #####
```

<!-- grade_id: correct_exercise_2, points: 6, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: correct_exercise_2 - possible points: 6

assert isinstance(cum_death, np.ndarray)
assert isinstance(cum_cases, np.ndarray)


##### DO NOT CHANGE #####
```

<!-- grade_id: correct_exercise_3_functions, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: correct_exercise_3_functions - possible points: 2

assert fun_1(0, 0, 0) == 0
assert fun_2(0, 0, 0, 0) == 0
assert fun_3(0, 0, 0, 0) == 0
assert fun_4(0, 0, 0, 0, 0) == 0


##### DO NOT CHANGE #####
```

<!-- grade_id: correct_exercise_3_fits_and_answers, points: 6, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: correct_exercise_3_fits_and_answers - possible points: 6

assert param_1 is not None
assert param_2 is not None
assert param_3 is not None
assert param_4 is not None
assert isinstance(answer_best_fun, int)
assert isinstance(answer_best_fun_14_days, int)


##### DO NOT CHANGE #####
```